Expression of Most Retrotransposons in Human Blood Correlates with Biological Aging

Author response:

The following is the authors’ response to the current reviews.

Reviewer #2 (Public Review):

The authors responded that they would lose statistical power by studying RTE subfamilies with limited microarray probes, which is a fair point. However, the suggested analysis could have been conducted using the RNA-seq data they explored in the second round of revision. Choosing not to leverage RNA-seq to increase the granularity of their analysis is a matter of choice. In my opinion, however, the authors could have acknowledged in the discussion that some smaller yet potentially influential RTE species may be masked by their global approach."

We will add one sentence addressing this in the Version of Record.

The following is the authors’ response to the original reviews.

We thank Reviewer #1 for their constructive comments.

Public Reviews:

Reviewer #1 (Public Review):

Tsai and Seymen et al. investigate associations between RTE expression and methylation and age and inflammation, using multiple public datasets. Compared to the previous round of review, the text of the manuscript has been polished and the phrasing of several findings has been made clearer and more precise. The authors also provided ample discussion to the prior reviewer comments in their rebuttal, including new analyses. All these changes are in the correct direction, however, I believe that part of the content of the rebuttal should be incorporated in the main text, for reasons that I will outline below.

Both reviewers found the reliance on microarray expression data to detract from the study. The authors argued that their choices are supported by existing publications which performed a similar quantification of TE expression using microarray data. It could still be argued that (as far as I can tell) Reichmann et al. used a substantially larger number of probes than this study, as a consequence of starting from different arrays, however, this is a minor point which the authors do not need to address. It is still undeniable that including the validation with RNA-seq data performed in the rebuttal would strengthen the manuscript. I especially believe that many readers would want to see this analysis be prominent in the manuscript, considering that both reviewers independently converged on the issue with microarray expression data. Personally, I would have included an RNA-seq dataset next to the microarray data in the main figures, however, I understand that this would require considerable restructuring and that placing RNAseq data besides array data might be misleading. Instead, I would ask that the authors include their rebuttal figures R1 and R2 as supplementary figures.

I would suggest introducing a new paragraph, between the section dedicated to expression data and the one dedicated to DNA methylation, mentioning the issues with microarray data (Some of which were mentioned by the reviewers and other which were mentioned by the authors in the discussion and introduction) to then introduce the validation with RNA-seq data.

We appreciate the reviewer’s understanding and detailed feedback. As suggested, Author response images 1 and 2 were added as supplementary figures to the manuscript, and one paragraph was added to the section investigating the correlation between RTE expression and chronological age. We have also added new descriptions to the introduction, discussion, and BAR analysis sections.

Author response image 3 is also a good addition and should be expanded to include the GTP and MESA study and possibly mentioned in the paragraph titled "RTE expression positively correlates with BAR gene signature scores except for SINEs."

We have updated Author response image 3 (now Author response image 1) to include GTP and MESA cohorts in the analysis. As shown in Author response image 1, except IFN-I and senescence scores on the MESA cohort that positively correlate with chronological ageing, the rest of the gene signatures display no positive correlation with chronological ageing.

Author response image 1 was originally created to separate the effect of chronological age and RTE expression on BAR gene signature scores. As it was meant to discriminate between BAR and chronological age, it doesn't provide additional information regarding the positive correlation between RTE expression and BAR gene signature that was not already present in the manuscript. Therefore, we did not add it to the manuscript.

Generalized linear models (GLM) analysis (BAR gene signature scores ~ RTE expression +chronological age).

For each RTE family, we separately performed GLM. Age (RTE family) indicates the chronological age when used in the design formula for that specific RTE family.

"In this study, we did not compare MESA with GTP etc. We have analysed each dataset separately based on the available data for that dataset. Therefore, sacrificing one analysis because of the lack of information from the other does not make sense. We would do that if we were after comparing different datasets. Moreover, the datasets are not comparable because they were collected from different types of blood samples."

Indeed, the datasets are not compared directly, but the associations between age, BER and TE expression for each dataset are plotted and discussed right next to each other. It is therefore natural to wonder if the differences between datasets are due to differences in the type of blood sample or if they are a consequence of the different probe sets. Using a common set of probes would help answer that question.

We understand that the reviewer is proposing a method to eliminate the possible causes of differences across datasets. However, incorporating such change would compromise the statistic power of MESA and GARP cohorts and also change our analysis structurally and digress from our main focus. Hence, we disagree to use the identical set of probes for all three cohorts.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Retrotransposons (RTEs) have been postulated to reactivate with age and contribute to aging through activated innate immune response and inflammation. Here, we analyzed the relationship between RTE expression and aging using published transcriptomic and methylomic datasets of human blood. Despite no observed correlation between RTEs activity and chronological age, the expression of most RTE classes and families except short interspersed nuclear elements (SINEs) correlated with biological age-associated gene signature scores. Strikingly, we found that the expression of SINEs was linked to upregulated DNA repair pathways in multiple cohorts. We also observed DNA hypomethylation with aging and significant increase in RTE expression level in hypomethylated RTEs except for SINEs. Additionally, our single-cell transcriptomic analysis suggested a role for plasma cells in aging mediated by RTEs. Altogether, our multi-omics analysis of large human cohorts highlights the role of RTEs in biological aging and suggests possible mechanisms and cell populations for future investigations.

Revised version: Retrotransposons (RTEs) have been postulated to reactivate with age and contribute to aging through activated innate immune response and inflammation. Here, we analyzed the relationship between RTE expression and aging using published transcriptomic and methylomic datasets of human blood. Despite no observed correlation between RTE activity and chronological age, the expression of most RTE classes and families except short interspersed nuclear elements (SINEs) correlated with biological age-associated gene signature scores. Strikingly, we found that the expression of SINEs was linked to upregulated DNA repair pathways in multiple cohorts. We also observed DNA hypomethylation with aging and significant increase in RTE expression level in hypomethylated RTEs except for SINEs. Additionally, our single-cell transcriptomic analysis suggested a role for plasma cells in aging mediated by RTEs. Altogether, our multi-omics analysis of large human cohorts highlights the role of RTEs in biological aging and suggests possible mechanisms and cell populations for future investigations.

Introduction

Transposable elements (TEs) are genetic elements that can move within the genome and are categorized into DNA transposons and retrotransposons (RTEs), which depend on cDNA intermediates to function. RTEs include three classes: endogenous retrovirus long terminal repeats (LTRs), and long and short interspersed nuclear elements, known as LINEs and SINEs, respectively. LINEs and SINEs employ target-primed reverse transcription for genome integration, with SINEs depending on the proteins encoded by LINEs1. Although most of the RTE sequences are dormant in the host genome, they continue to play vital roles in human evolution and physiology2,3.

RTEs are silenced through heterochromatinization and DNA methylation in the early developmental stages as part of the host defense mechanism3,4. However, increased RTE activity that occurs with aging can lead to genome instability and activation of DNA damage pathways4–7. Additionally, accumulation of cytoplasmic RTE cDNAs detected in aging organisms and senescent cells can activate type I interferon (IFN-I) response and inflammation3,4,8–10, that contributes to inflammaging11. Furthermore, chromatin remodeling and cellular senescence also contribute partially to the RTE reactivation10. Cellular senescence is a non-proliferative state elicited by stress factors such as DNA damage and is closely intertwined with inflammation and aging12. Through the expression of senescence associated secretory phenotype (SASP), senescent cells promote inflammation and senescence in other cells, including immune cells, leading to a compromised immune system and a vicious cycle of inflammation13. Understanding the relationship between RTE expression and SASP, cellular senescence, and inflammaging could offer important insights into strategies against aging and age-related diseases (ARDs).

Earlier version: Recent studies have documented the relationship between RTE activation and biological age-related (BAR) events, but comprehensive and large-scale studies are lacking, mainly due to the paucity of RNA-seq data for large non-cancerous human cohorts. Most of the existing transcriptomics datasets are microarray-based, where the computational methods to analyze repetitive elements have not been sufficiently developed. Overlapping Illumina microarray expression and methylation probe locations to RTE locations in RepeatMasker14, we were able to identify enough probes to calculate the expression of RTE classes and families. Building on this methodology, we explored how RTE expression contributes to biological aging using publicly available transcriptomics microarray data derived from human blood samples. More specifically, we first investigated if RTE expression was correlated with chronological age, and then we analyzed the relationship between RTE and BAR events including cellular senescence, inflammation, and IFN-I response with published gene signatures. Using microarray methylomic data, we also investigated the DNA methylation level of RTEs in blood samples of multiple non-cancerous human cohorts and examined the relation between DNA methylation and RTE expression and aging. Furthermore, with annotated single-cell transcriptomic data of the peripheral blood mononuclear cells (PBMCs) of 21 healthy human samples from a recent study of aging immunity15, we identified cells that could be implicated in the process of RTE reactivation and aging. Lastly, the single-cell transcriptomic data of supercentenarians16 was analyzed to examine the interplay between RTE expression and aging as a factor of aging and longevity (Fig. 1).

Revised version: Recent studies have documented the relationship between RTE activation and biological age-related (BAR) events, but comprehensive and large-scale studies are lacking, mainly due to the paucity of RNA-seq data for large non-cancerous human cohorts. Most of the existing transcriptomics datasets are microarray-based, where the computational methods to analyze repetitive elements have not been sufficiently developed. By overlapping Illumina microarray expression and methylation probe locations to RTE locations in RepeatMasker14, we were able to identify sufficient number of probes to calculate the expression and methylation levels of RTE classes and families. Building on this methodology, we explored how RTE expression contributes to biological aging using publicly available transcriptomics microarray data derived from human blood samples. More specifically, we first investigated if RTE expression was correlated with chronological age, and then we analyzed the relationship between RTE and BAR events including cellular senescence, inflammation, and IFN-I response with published gene signatures. Using microarray methylomic data, we also investigated the DNA methylation level of RTEs in blood samples of multiple non-cancerous human cohorts and examined the relation between DNA methylation and RTE expression and aging. Furthermore, with annotated single-cell transcriptomic data of the peripheral blood mononuclear cells (PBMCs) of 21 healthy human samples from a recent study of aging immunity15, we identified cells that could be implicated in the process of RTE reactivation and aging. We also validated our microarray result with a publicly available bulk RNA-seq PBMC data from healthy individuals that was recently published in a new study of aging16. Lastly, the single-cell transcriptomic data of supercentenarians17 was analyzed to examine the interplay between RTE expression and aging as a factor of aging and longevity (Fig. 1).

Conceptual framework and the study design.

We collected published datasets of human blood samples for gene expression, DNA methylation, and single-cell transcriptomic data. The analysis aimed to study the relation between the expression and DNA methylation of RTEs versus chronological and biological aging in large human cohorts. The single-cell transcriptomic datasets were employed for cell type-specific analysis of RTEs in PBMC to identify the relation between RTE expression and aging events for annotated cell types within old versus young PBMC samples.

Results

Analyzing DNA methylation and expression levels of RTEs using microarray data

Earlier version: We collected three published microarray datasets from large-scale human studies, including the peripheral monocyte samples from Multi-Ethnic Study of Atherosclerosis17 (MESA, aged 44 to 83, n=1202), the whole blood (WB) samples from Grady Trauma Project18,19 (GTP, aged 15 to 77, n=359), and PBMC samples from Genetics, Osteoarthritis and Progression20 (GARP, aged 43 to 79, n=139), which are all non-cancerous human samples (Supplementary Table 1 and Supplementary Information). The studies were conducted using either Illumina HumanHT-12 V3 or Illumina HumanHT-12 V4 expression microarray kits. After comparing the probes of the two microarray versions, we adopted the more comprehensive probe list of V4, which contains the full intersection of the two lists. To quantify RTE expression, we mapped the microarray probe locations to RTE locations in RepeatMasker to extract the list of noncoding (intergenic or intronic) probes that cover the RTE regions. We included three main RTE classes: (1) the LINE class, which encompasses the L1 and L2 families; (2) the SINE class, with Alu and MIR as the two main families; and (3) the LTR class which comprises the ERV1, ERVL, ERVL-MaLR, and ERVK families. Most of the RTE-covering probes available on Illumina HumanHT-12 V4 are present in MESA and GARP, while fewer are available in GTP (Supplementary Fig. 1a, Supplementary Tables 2-4).

Revised version: We collected three published microarray datasets from large-scale human studies, including the peripheral monocyte samples from Multi-Ethnic Study of Atherosclerosis18 (MESA, aged 44 to 83, n=1202), the whole blood (WB) samples from Grady Trauma Project19,20 (GTP, aged 15 to 77, n=359), and PBMC samples from Genetics, Osteoarthritis and Progression21 (GARP, aged 43 to 79, n=139), which are all non-cancerous human samples (Supplementary Table 1 and Supplementary Information). The studies were conducted using either Illumina HumanHT-12 V3 or Illumina HumanHT-12 V4 expression microarray kits. After comparing the probes of the two microarray versions, we adopted the more comprehensive probe list of V4, which contains the full intersection of the two lists. To quantify RTE expression, we mapped the microarray probe locations to RTE locations in RepeatMasker to extract the list of noncoding (intergenic or intronic) probes that cover the RTE regions. We included three main RTE classes: (1) the LINE class, which encompasses the L1 and L2 families; (2) the SINE class, with Alu and MIR as the two main families; and (3) the LTR class which comprises the ERV1, ERVL, ERVL-MaLR, and ERVK families. Most of the RTE-covering probes available on Illumina HumanHT-12 V4 are present in MESA and GARP, while fewer are available in GTP (Supplementary Fig. 1a, Supplementary Tables 2-4).

Earlier version: Four methylation datasets were analysed, including MESA, Swedish Adoption/Twin Study of Aging21 (SATSA, aged 48 to 98, n=1072), Brisbane Systems Genetics Study22 (BSGS, aged 10 to 75, n=862), and Genome-wide Methylation Profiles Reveal Quantitative Views of Human Aging Rates23 (GMPWAR, aged 19 to 101, n=656). SATSA, BSGS, and GMPWAR include genome-wide DNA methylation of WB samples produced by Illumina Infinium 450k array (Supplementary Table 1 and Supplementary Information). The DNA methylation probes from the Illumina Infinium 450k array kit were aligned to the locations in RepeatMasker to identify the probes overlapping the RTE regions (Supplementary Table 5). More than 90% of the RTE-covering probes are present in MESA and GMPWAR, while fewer, but more than half, are available in SATA and BSGS datasets (Supplementary Fig. 1b).

Revised version: Four methylation datasets were analysed, including MESA, Swedish Adoption/Twin Study of Aging22 (SATSA, aged 48 to 98, n=1072), Brisbane Systems Genetics Study23 (BSGS, aged 10 to 75, n=862), and Genome-wide Methylation Profiles Reveal Quantitative Views of Human Aging Rates24 (GMPWAR, aged 19 to 101, n=656). SATSA, BSGS, and GMPWAR include genome-wide DNA methylation of WB samples produced by Illumina Infinium 450k array (Supplementary Table 1 and Supplementary Information). The DNA methylation probes from the Illumina Infinium 450k array kit were aligned to the locations in RepeatMasker to identify the probes overlapping the RTE regions (Supplementary Table 5). More than 90% of the RTE-covering probes are present in MESA and GMPWAR, while fewer, but more than half, are available in SATA and BSGS datasets (Supplementary Fig. 1b).

The chronological age is not linked with RTE expression

Firstly, we examined the relationship between RTE expression and chronological age in the MESA, GARP, and GTP cohorts. The ages of the individuals enrolled in these studies range from 15 to 83 years, with the mean being 70.2 (MESA), 60.1 (GAPR), and 42.5 (GTP) (Supplementary Table 1). Strikingly, no significant relationship (P value < 0.05) is found between chronological age and expressions of the three RTE classes across the datasets (Fig. 2a). Similarly, apart from a weak correlation observed between a few RTE families and chronological age, we could not identify any strong relation between chronological age and expression of RTE families across the three cohorts (Supplementary Fig. 2a).

Correlation analysis between RTE expression, chronological age, and age-associated gene signature scores.

a, No correlation between RTE expression and chronological age versus positive correlations between BAR gene signature scores and LINE and LTR expressions. Pair-wise correlation coefficients were calculated between the expression of different RTE classes (LTR, LINE & SINE) and chronological age and six BAR gene signature scores in monocytes (MESA), PBMCs (GARP), and the WB (GTP). b, Scatter plots displaying a positive correlation between LINE and LTR expressions and inflammaging, SASP, and senescence gene signature scores in PBMCs. c, Different families of RTEs were divided into two major groups based on their correlation and inverse correlation with BAR gene signature scores in PBMC samples. d, Correlation matrix depicting all pair-wise combinations to identify the correlation between chronological age, RTE family expressions, and six age-associated signature scores in PBMCs. ** P ≤ 0.01, *** P ≤ 0.001, Pearson’s correlation. MESA, n=1202; GARP, n=139; GTP, n=359.

Earlier version: A similar correlation analysis was carried out in a scRNA-seq data of the PBMCs of 21 non-obese healthy men annotated into 25 cell types by Mogilenko et al.15 (Supplementary Table 6 and online methods). When comparing the RTE expression of the young (aged 25 to 29, n=11) group with the old (aged 62 to 70, n=10) group, we could not find any significant difference or trend in any cell types between the two groups (Supplementary Figs. 3-5). Overall, RTE expression did not correlate with chronological aging across our microarray and scRNA-seq datasets.

Revised version: A similar correlation analysis was carried out in a scRNA-seq data of the PBMCs of 21 non-obese healthy men annotated into 25 cell types by Mogilenko et al.15 (Supplementary Table 6 and online methods). When comparing the RTE expression of the young (aged 25 to 29, n=11) group with the old (aged 62 to 70, n=10) group, we could not find any significant difference or trend in any cell types between the two groups (Supplementary Figs. 3-5).

Added in the revised version: To investigate whether our observation is consistent in whole transcriptomic data, we conducted correlation analysis on a set of RNA-seq data on healthy human PBMC16 (aged 20 to 74, n=159). We did not observe any statistically significant correlation between chronological age and expression of RTE classes or families (Supplementary Fig. 6). Overall, RTE expression did not correlate with chronological aging across the microarray, scRNA-seq, and RNA-seq datasets.

RTE expression positively correlates with BAR gene signature scores except for SINEs

Earlier version: In mammals, aging is characterized by an intricate network of inflammation, IFN-I signaling, and cellular senescence3. To measure the level of biological aging from the transcriptomic data, we focused on six gene sets retrieved from widely cited studies: IFN-I24, inflammatory cytokines25, inflammatory chemokines26, inflammaging27, SASP28, and cellular senescence29 (Supplementary Table 7). The scores of these six age-associated gene sets were calculated for each individual in the MESA, GTP, and GARP cohorts by applying the singscore package in Bioconductor30 (online methods). Except for GTP, in which 12 genes were missing from the gene sets in total (Supplementary Table 8), MESA and GARP microarray datasets had the expression of all the genes in the gene sets.

Revised version: In mammals, aging is characterized by an intricate network of inflammation, IFN-I signaling, and cellular senescence3. To measure the level of biological aging from the transcriptomic data, we focused on six gene sets retrieved from widely cited studies: IFN-I25, inflammatory cytokines26, inflammatory chemokines27, inflammaging28, SASP29, and cellular senescence30 (Supplementary Table 7). The scores of these six age-associated gene sets were calculated for each individual in the MESA, GTP, and GARP cohorts by applying the singscore package in Bioconductor31 (online methods). Except for GTP, in which 12 genes were missing from the gene sets in total (Supplementary Table 8), MESA and GARP microarray datasets had the expression of all the genes in the gene sets.

Earlier version: Then, we analyzed the relation between RTE expression and the six BAR gene signature scores. Positive correlations were observed between the expressions of LINEs and LTRs and BAR gene signature scores across three datasets, while inverse correlation was present between the expression of SINEs and these scores in the PBMC samples from the GARP cohort (Fig. 2a and b). The discrepant pattern of LTRs and LINEs versus SINEs in PBMC is more pronounced in the correlation analysis between the expression of RTE families and the BAR gene sets (Fig. 2c and d). Interestingly, the expressive patterns of RTE families are divided into two groups that also reflect their positive or inverse correlations with BAR signature scores in the GARP cohort. In the GARP cohort, the strongest positive correlation with BAR scores was observed in L1, ERVK, and ERVL. By contrast, the expressions of Alu and MIR of SINE, L2, and ERV1 were inversely correlated with BAR signature scores in the GARP cohort (Fig. 2c and d). This pattern was also observed in the MESA peripheral monocyte samples for L1, ERVL, and ERVK, but not other RTE families (Supplementary Fig.2a and c). The WB GTP cohort showed divergent result to the other two cohorts perhaps because of the absence of certain genes in the six BAR gene sets or lower number of TE probes in GTP (Supplementary Fig. 1b, Fig.2a-b and Supplementary Table 8).

Revised version: We analyzed the relation between RTE expression and the six BAR gene signature scores. Positive correlations were observed between the expressions of LINEs and LTRs and BAR gene signature scores across three datasets, while inverse correlation was present between the expression of SINEs and these scores in the PBMC samples from the GARP cohort (Fig. 2a and b). The discrepant pattern of LTRs and LINEs versus SINEs in PBMC is more pronounced in the correlation analysis between the expression of RTE families and the BAR gene sets (Fig. 2c and d). Interestingly, the expressive patterns of RTE families are divided into two groups that also reflect their positive or inverse correlations with BAR signature scores in the GARP cohort. In the GARP cohort, the strongest positive correlation with BAR scores was observed in L1, ERVK, and ERVL. By contrast, the expressions of Alu and MIR of SINE, L2, and ERV1 were inversely correlated with BAR signature scores in the GARP cohort (Fig. 2c and d). This pattern was also observed in the MESA peripheral monocyte samples for L1, ERVL, and ERVK, but not other RTE families (Supplementary Fig.2a and c). The WB GTP cohort showed divergent result to the other two cohorts perhaps because of the absence of certain genes in the six BAR gene sets or lower number of TE probes in GTP (Supplementary Fig. 1b, Fig.2a-b and Supplementary Table 8).

Earlier version: To study the association between RTE expression and BAR gene signatures, we also split the samples in each cohort based on quartiles of RTE expression into high, low, and medium (top 25%, bottom 25%, and middle 50%) groups and compared the BAR gene signature scores across these three groups. Across the datasets, there is an overall trend of increased scores of BAR gene signatures in high versus low LINE and LTR expression groups (Supplementary Fig. 6). In fact, amongst all gene sets, we found the most significant increase for SASP and senescence and to a lesser extent for inflammaging signatures. We observed higher SASP, inflammatory cytokine, senescence, and inflammaging in the peripheral monocytes (MESA) in high SINE- and Alu-expressing groups (Supplementary Fig. 7). However, in the PBMCs samples of the GARP cohorts, groups of high SINE and Alu expressions demonstrate lower expressions of SASP, inflammatory chemokine, senescence, and inflammaging, that reflects the inverse correlation we observed before (Fig. 2c). There are no specific trends for SINE expression in the WB samples from the GTP cohort (Supplementary Fig. 7). Taken together, SINEs display different pattern than LINEs and LTRs in terms of the association between their expressions and BAR gene signatures. This prompted us to conduct a gene set variation analysis (GSVA)31 to identify the pathways that are regulated in the high versus low RTE expression groups in the three cohorts (online methods).

Revised version: To study the association between RTE expression and BAR gene signatures, we also split the samples in each cohort based on quartiles of RTE expression into high, low, and medium (top 25%, bottom 25%, and middle 50%) groups and compared the BAR gene signature scores across these three groups. Across the datasets, there is an overall trend of increased scores of BAR gene signatures in high versus low LINE and LTR expression groups (Supplementary Fig. 8). In fact, amongst all gene sets, we found the most significant increase for SASP and senescence and to a lesser extent for inflammaging signatures. We observed higher SASP, inflammatory cytokine, senescence, and inflammaging in the peripheral monocytes (MESA) in high SINE- and Alu-expressing groups (Supplementary Fig. 9). However, in the PBMC samples of the GARP cohorts, groups of high SINE and Alu expressions demonstrate lower expressions of SASP, inflammatory chemokine, senescence, and inflammaging, that reflects the inverse correlation we observed before (Fig. 2c). Interestingly, when we examined RNA-seq PBMC data, we found that there is a significant correlation between IFN-I signature score and expression of LINE and LTR classes and most of their families (P-val < 0.01), but such a significant correlation did not exist between IFN-I score and SINE class and Alu family (Supplementary Fig. 7). There are also no specific trends for SINE expression in the WB samples from the GTP cohort (Supplementary Fig. 9). Taken together, SINEs display different pattern than LINEs and LTRs in terms of the association between their expressions and BAR gene signatures. This prompted us to conduct a gene set variation analysis (GSVA)32 to identify the pathways that are regulated in the high versus low RTE expression groups in the three cohorts (online methods).

Elevated SINE expression is linked with the upregulation of the DNA repair pathways, while elevated LINE and LTR expressions are associated with the inflammatory response

Earlier version: We conducted GSVA for high versus low expression groups of RTE classes and families to identify the pathways that might be affected by RTE expression. We specifically focused on the pathways (gene sets) related to the inflammatory and DNA repair responses retrieved from the Molecular Signatures Database (MSigDB)32 with “inflammatory” and “dna_repair” as search keywords.

Revised version: We conducted GSVA for high versus low expression groups of RTE classes and families to identify the pathways that might be affected by RTE expression. We specifically focused on the pathways (gene sets) related to the inflammatory and DNA repair responses retrieved from the Molecular Signatures Database (MSigDB)33 with “inflammatory” and “dna_repair” as search keywords.

Earlier version: We found that a significant fraction of gene sets related to DNA repair were downregulated in the group of monocyte samples (MESA cohort) highly expressing LINE and LTR classes but upregulated in the monocyte and PBMC samples (MESA and GARP cohorts, respectively) highly expressing SINEs (Fig. 3a). At the family level, there is a unique increase in DNA repair activity in the monocytes of the cohort with high Alu expression. In the PBMC samples of the GARP dataset, DNA repair pathways are upregulated in groups highly expressing L2, ERV1, Alu, and MIR, while most of the other RTE classes demonstrate decreased DNA repair activity.

Revised version: We found that a significant fraction of gene sets related to DNA repair were downregulated in the group of monocyte samples (MESA cohort) highly expressing LINE and LTR classes but upregulated in the monocyte and PBMC samples (MESA and GARP cohorts, respectively) highly expressing SINEs (Fig. 3a). At the family level, there is a unique increase in DNA repair activity in the monocytes of the cohort with high Alu expression. In the PBMC samples of the GARP dataset, DNA repair pathways are upregulated in groups highly expressing L2, ERV1, Alu, and MIR, while most of the other RTE classes demonstrate decreased DNA repair activity. In contrast, the inflammatory response is upregulated in the monocyte and PBMC samples of groups expressing LINE and LTR classes and families at high levels, and such upregulation is less common in SINE (Fig. 3a). These observations provide insight into the perpetuating contrast between SINE versus LINE and LTR observed in our analysis, as SINE might contribute to aging via genome instability instead of inflammation that is associated with LINE and LTR activities. Overall, our observations are consistent across MESA and GARP cohorts but less so in the GTP cohort (Fig. 3b), which might be attributed to less RTE-covering probes and the missing genes from the BAR gene signatures. (Supplementary Fig. 1a, Supplementary Table 8).

Upregulation of DNA repair vs inflammatory responses for samples with high expression of SINE vs LTR and LINE in the MESA and GARP cohorts.

a, GSVA demonstrates increased activity of DNA repair pathways in the group of samples with high vs low SINE expression in the MESA and GARP cohorts. In contrast, the inflammatory response is upregulated in the sample groups highly expressing LINE and LTR classes and families in the MESA and GARP cohorts. The samples in each cohort were divided into low (1st quartile), medium (2nd and 3rd quartile), and high (4th quartile) expression groups based on the expression of RTE classes or families. GSVA was applied on high vs low groups for each class and family of RTEs. The threshold for differential expression is set at |logFC| > 0.1 and P < 0.05 (online methods). b, The Radar plot shows the difference between the number of upregulated versus downregulated gene sets related to DNA repair and inflammatory response in each cohort. While high expression of SINE and Alu is significantly associated with high number of up-regulated DNA-repair gene sets, LINE and L1 expressions are highly related to high number of activated gene sets related to inflammatory response in the MESA and GARP cohorts. This result is not highly supported by the GTP cohort, more likely due to the low number of probes in this cohort.

Removed in the revised version: In contrast, the inflammatory response is upregulated in the monocyte and PBMC samples of groups expressing LINE and LTR classes and families at high levels, and such upregulation is less common in SINE (Fig. 3a). These observations provide insight into the perpetuating contrast between SINE versus LINE and LTR observed in our analysis, as SINE might contribute to aging via genome instability instead of inflammation that is associated with LINE and LTR activities. Overall, our observations are consistent across MESA and GARP cohorts but less so in the GTP cohort (Fig. 3b), which might be attributed to less RTE-covering probes and the missing genes from the BAR gene signatures. (Supplementary Fig. 1a, Supplementary Table 8).

DNA methylation levels of RTEs inversely correlated with the chronological age and the RTE expression except for SINE expression

Earlier version: The demethylation of RTEs is suggested to contribute to aging3,5,7. To examine the effect of RTE demethylation on chronological aging, we analyzed DNA methylation data from multiple cohorts including monocytes from the MESA cohort and the WB data from GMPWAR, BSGS, SATSA. Our analysis showed an inverse correlation between the methylation level of RTEs and chronological age across all RTE classes and families in the WB and monocytes (Fig. 4a-b and Supplementary Fig. 8). This observed trend is in line with previous reports of demethylation of RTE with age in gene-poor regions5,21,23.

Revised version: The demethylation of RTEs is suggested to contribute to aging3,5,7. To examine the effect of RTE demethylation on chronological aging, we analyzed DNA methylation data from multiple cohorts including monocytes from the MESA cohort and the WB data from GMPWAR, BSGS, SATSA. Our analysis showed an inverse correlation between the methylation level of RTEs and chronological age across all RTE classes and families in the WB and monocytes (Fig. 4a-b and Supplementary Fig. 10). This observed trend is in line with previous reports of demethylation of RTE with age in gene-poor regions5,22,24.

Inverse correlation of DNA methylation levels of RTEs with the chronological age and the RTE expression except for SINE expression.

a, b, Methylation levels of RTE classes inversely correlated with chronological age in monocyte (MESA) and WB (BSGS, SATSA, and GMPWAR) samples. Satellite DNA was included as a control group. c, Methylation levels versus low (1st quartile), medium (2nd and 3rd quartile), and high (4th quartile) expressions of RTE classes in monocytes (MESA). Wilcoxon test; ns: not significant. d, Correlation matrix for RTE expressions and methylation levels, and chronological age. * P ≤ 0.05, **** P ≤ 0.0001, Pearson’s correlation. MESA, n=1202; BSGS, n=614; GMPWAR, n=656; SATSA, n=1072.

Earlier version: To elucidate the relationship between the expression and methylation level of RTEs, we divided the MESA samples, with matched RTE DNA methylation and expression, into three groups based on the level of expression of RTEs in different classes and families: low (1st quartile), medium (2nd and 3rd quartile), and high (4th quartile) expression groups of monocyte samples from the MESA cohort. The DNA methylation levels of LINEs and LTRs significantly decreased in groups of high LINE and LTR expressions, respectively, but SINEs and the satellite repeats (as negative control) did not display the same pattern (Fig. 4c). LINE families, MIR, and LTR families except ERVK displayed lower levels of methylation in higher expression groups, while such pattern was absent in Alu, CR1, and ERVl (Supplementary Fig. 9). Although the DNA methylation level of SINEs positively correlated with the methylation levels of LINEs and LTRs, the expression of SINEs was not significantly correlated with the expression levels of other RTE classes (Fig. 4d). Overall, we found a distinct pattern for SINEs versus LTRs and LINEs by exploring a relationship between DNA methylation and expression levels of RTEs in monocytes, consistent with our other findings showing different patterns for SINEs versus LTRs and LINEs in PBMCs (see above). This result also suggests a dispensable role of DNA methylation for SINE silencing in monocytes, which is in line with previous studies showing SINEs are not derepressed by deletion of DNA methyltransferases or treatment with DNA demethylating agents and are primarily regulated by histone modifications33.

Revised version: To elucidate the relationship between the expression and methylation level of RTEs, we divided the MESA samples, with matched RTE DNA methylation and expression, into three groups based on the level of expression of RTEs in different classes and families: low (1st quartile), medium (2nd and 3rd quartile), and high (4th quartile) expression groups of monocyte samples from the MESA cohort. The DNA methylation levels of LINEs and LTRs significantly decreased in groups of high LINE and LTR expressions, respectively, but SINEs and the satellite repeats (as negative control) did not display the same pattern (Fig. 4c). LINE families, MIR, and LTR families except ERVK displayed lower levels of methylation in higher expression groups, while such pattern was absent in Alu, CR1, and ERVl (Supplementary Fig. 11). Although the DNA methylation level of SINEs positively correlated with the methylation levels of LINEs and LTRs, the expression of SINEs was not significantly correlated with the expression levels of other RTE classes (Fig. 4d). Overall, we found a distinct pattern for SINEs versus LTRs and LINEs by exploring a relationship between DNA methylation and expression levels of RTEs in monocytes, consistent with our other findings showing different patterns for SINEs versus LTRs and LINEs in PBMCs (see above). This result also suggests a dispensable role of DNA methylation for SINE silencing in monocytes, which is in line with previous studies showing SINEs are not derepressed by deletion of DNA methyltransferases or treatment with DNA demethylating agents and are primarily regulated by histone modifications34.

Elevated level of RTEs in plasma cells of healthy PBMC samples is associated with the high SASP and inflammaging gene signature scores

Earlier version: As described above, we did not identify any association between RTE expression and chronological aging for any human cohorts we analyzed. However, the positive correlation between LINE and LTR expressions with BAR gene signature scores in the PBMC samples of the GARP cohort prompted us to investigate PBMC samples in more detail using PBMC scRNAseq data of 21 samples annotated into 25 cell types by Mogilenko et al.15(online methods and Supplementary Table 6). The scRNA-seq samples were divided based on low and high BAR gene signature scores to compare against the expression of RTE classes for each cell type (online methods). Among the 25 annotated cell types, plasma cells consistently demonstrate significantly elevated RTE expression in cells of high SASP expression that is not observed in the other cell types (Fig. 5a and Supplementary Figs. 10-12). Increased LINE and SINE expressions were also observed in the plasma cells of the samples with high inflammatory chemokines and inflammaging gene signature scores (Fig. 5b and c). Although not statistically significant, elevated RTE expression was also seen in the groups with high inflammatory cytokines and senescence but not IFN-I scores (Fig. 5d-f). Overall, this finding indicates that plasma cells might play an important role in the process in which reactivation of RTEs regulates aging.

Revised version: As described above, we did not identify any association between RTE expression and chronological aging for any human cohorts we analyzed. However, the positive correlation between LINE and LTR expressions with BAR gene signature scores in the PBMC samples of the GARP cohort prompted us to investigate PBMC samples in more detail using PBMC scRNAseq data of 21 samples annotated into 25 cell types by Mogilenko et al.15(online methods and Supplementary Table 6). The scRNA-seq samples were divided based on low and high BAR gene signature scores to compare against the expression of RTE classes for each cell type (online methods). Among the 25 annotated cell types, plasma cells consistently demonstrate significantly elevated RTE expression in cells of high SASP expression that is not observed in the other cell types (Fig. 5a and Supplementary Figs. 12-14). Increased LINE and SINE expressions were also observed in the plasma cells of the samples with high inflammatory chemokines and inflammaging gene signature scores (Fig. 5b and c). Although not statistically significant, elevated RTE expression was also seen in the groups with high inflammatory cytokines and senescence but not IFN-I scores (Fig. 5d-f). Overall, this finding indicates that plasma cells might play an important role in the process in which reactivation of RTEs regulates aging.

Cell type-specific analysis of RTE expression vs BAR gene signature scores in two PBMC scRNA-seq cohorts.

a-f, Unique increased expression of RTEs in the Plasma B cells of the samples with high BAR gene signature scores indicates the potential role of plasma B cells in aging. Wilcoxon test. n = 21. g, Decreased expressions of RTE classes and increased BAR gene signature scores in multiple annotated cell types obtained from the PBMCs of supercentenarians compared to ordinary elderlies as control. Supercentenarians, n=7; control, n=5, age 50 to 80. Wilcoxon test was applied to identify the significant changes. NK, Natural killer cell; BC, B-cell; TC1, T-cell 1; TC2, T-cell 2; M14, CD14+ monocyte; M16, CD16+ monocyte; EC, Erythrocytes; MKI, MKI67+ proliferating cell; DC, Dendritic cell; MGK, Megakaryocyte.

Downregulation of RTE expression in supercentenarians versus normal aged cases despite high BAR signature scores in NK and T cells of supercentenarians

Earlier version: Supercentenarians who have reached 110 years of age are an excellent resource for investigating healthy aging. We hypothesized that supercentenarians should express RTEs less than normal-aged people because RTE expression, particularly LINE and LTR expressions, can activate the inflammatory pathways (Fig. 3), which would be generally harmful to healthy aging. To understand whether there is any RTE expression change in supercentenarians versus normal-aged people and identify the relationship to BAR gene signatures, we compared the scRNA-seq data of the PBMCs of seven supercentenarians (over 110 years) to five non-centenarian (50 to 80 years) controls16. We collected the annotated cells of these twelve cases as a Seurat object for which the cells were annotated into ten different cell types. Compared with the control group, the supercentenarians demonstrated decreased RTE expression in the natural killer (NK) cells, B-cells, T-cells, monocytes, and dendritic cells (Fig. 5g and Supplementary Figs. 13). Both T cell types showed decreased LINE expression, while the noncytotoxic cluster, TC1, showed slightly decreased LTR expression, and the expanding cytotoxic T cells, TC2, showed decreased SINE expression. On the other hand, the cytotoxic T cells also demonstrated an increased level of inflammatory cytokines and inflammaging gene signature scores. In addition, the NK cells displayed higher senescence and SASP scores in the supercentenarians while demonstrating decreased LINE and SINE expressions (Fig. 5g and Supplementary Figs. 14-15). Overall, this result reveals inflammatory activity in cytotoxic T cells and increased senescence in the NK cells that is not associated with the activation of RTEs in supercentenarians. By contrast, the expression of RTEs is significantly reduced in most of the immune cell types in supercentenarians.

Revised version: Supercentenarians who have reached 110 years of age are an excellent resource for investigating healthy aging. We hypothesized that supercentenarians should express RTEs less than normal-aged people because RTE expression, particularly LINE and LTR expressions, can activate the inflammatory pathways (Fig. 3), which would be generally harmful to healthy aging. To understand whether there is any RTE expression change in supercentenarians versus normal-aged people and identify the relationship to BAR gene signatures, we compared the scRNA-seq data of the PBMCs of seven supercentenarians (over 110 years) to five non-centenarian (50 to 80 years) controls17. We collected the annotated cells of these twelve cases as a Seurat object for which the cells were annotated into ten different cell types. Compared with the control group, the supercentenarians demonstrated decreased RTE expression in the natural killer (NK) cells, B-cells, T-cells, monocytes, and dendritic cells (Fig. 5g and Supplementary Fig. 15). Both T cell types showed decreased LINE expression, while the noncytotoxic cluster, TC1, showed slightly decreased LTR expression, and the expanding cytotoxic T cells, TC2, showed decreased SINE expression. On the other hand, the cytotoxic T cells also demonstrated an increased level of inflammatory cytokines and inflammaging gene signature scores. In addition, the NK cells displayed higher senescence and SASP scores in the supercentenarians while demonstrating decreased LINE and SINE expressions (Fig. 5g and Supplementary Figs. 16-17). Overall, this result reveals inflammatory activity in cytotoxic T cells and increased senescence in the NK cells that is not associated with the activation of RTEs in supercentenarians. By contrast, the expression of RTEs is significantly reduced in most of the immune cell types in supercentenarians.

Discussion

Earlier version: The role of RTEs in aging has been studied extensively but not in a systematic manner with large, non-cancerous human cohorts across RTE categories before this study. Current understanding of the relationship between RTEs and aging are largely conducted in vitro or in model organisms7,9,34 that might not apply to human due to intraspecies differences35. Moreover, the link of RTEs to aging in non-cancerous human cohorts has not been studied systematically in large cohorts.

Revised version: The role of RTEs in aging has been studied extensively but not in a systematic manner with large, non-cancerous human cohorts across RTE categories before this study. Current understanding of the relationship between RTEs and aging are largely conducted in vitro or in model organisms7,9,35 that might not apply to human due to intraspecies differences36. Moreover, the link of RTEs to aging in non-cancerous human cohorts has not been studied systematically in large cohorts.

Earlier version: We combined bulk and single-cell transcriptomic data to examine the relationship between RTE expression versus chronological and biological aging, demonstrating that chronological aging is not significantly linked with RTE reactivation. In this process, we discovered that LTR and LINE expressions are positively correlated with the BAR gene signature scores. However, in contrast, the SINE expression demonstrates an inverse correlation with BAR scores in the PBMC samples and no significant relationship in the WB. We suspect the absence of a positive correlation between BAR scores and SINE expression might be linked to the same discrepancy observed in the GSVA analysis, in which DNA repair pathways are upregulated in the groups of cases with high SINE and Alu expressions. This result suggests that reactivation of SINE, particularly Alu, might lead to DNA damage, a hallmark of aging. Although SINEs compose up to 13% of the human genome, and Alu elements make up the largest family of human mobile elements, Alu is not studied as extensively as L1 from the LINE class1. Therefore, this novel discovery requires further study to confirm the connection and elucidate the mechanism by which DNA repair might be linked to the derepression of SINE and Alu expressions. In contrast, the groups of cases with high LINE and LTR expressions have upregulation of the pathways related to the inflammatory response, which is in line with the previous findings from human and model organisms3,34.

Revised version: We combined bulk and single-cell transcriptomic data to examine the relationship between RTE expression versus chronological and biological aging, demonstrating that chronological aging is not significantly linked with RTE reactivation. In this process, we discovered that LTR and LINE expressions are positively correlated with the BAR gene signature scores. However, in contrast, the SINE expression demonstrates an inverse correlation with BAR scores in the PBMC samples and no significant relationship in the WB. We suspect the absence of a positive correlation between BAR scores and SINE expression might be linked to the same discrepancy observed in the GSVA analysis, in which DNA repair pathways are upregulated in the groups of cases with high SINE and Alu expressions. This result suggests that reactivation of SINE, particularly Alu, might lead to DNA damage, a hallmark of aging. Although SINEs compose up to 13% of the human genome, and Alu elements make up the largest family of human mobile elements, Alu is not studied as extensively as L1 from the LINE class1.

Added in the revised version: Therefore, this novel discovery requires further study to confirm the connection and elucidate the mechanism by which DNA repair might be linked to the derepression of SINE and Alu expressions. In contrast, the groups of cases with high LINE and LTR expressions have upregulation of the pathways related to the inflammatory response, which is in line with the previous findings from human and model organisms3,35.

Earlier version: The mechanisms of aging and age-related diseases (ARDs) converge on inflammaging36. Previous studies have shown RTE’s pathogenetic role in ARDs including neurodegenerative diseases through elevation of DNA damage and genome instability37–39. SINE expression inversely correlated with the BAR signature scores in the PBMCs of osteoarthritis patients, while a positive correlation is observed in monocytes. Moreover, a discrepant upregulation of DNA repair pathways is seen in the osteoarthritis cohorts expressing higher level L2, ERV1, Alu, and MIR. With osteoarthritis being a typical ARD, our analysis provides insight into the relationship between aging and RTEs in an ARD state in comparison with that of non-ARD individuals.

Revised version: The mechanisms of aging and age-related diseases (ARDs) converge on inflammaging37. Previous studies have shown RTE’s pathogenetic role in ARDs including neurodegenerative diseases through elevation of DNA damage and genome instability38–40. SINE expression inversely correlated with the BAR signature scores in the PBMCs of osteoarthritis patients, while a positive correlation is observed in monocytes. Moreover, a discrepant upregulation of DNA repair pathways is seen in the osteoarthritis cohorts expressing higher level L2, ERV1, Alu, and MIR. With osteoarthritis being a typical ARD, our analysis provides insight into the relationship between aging and RTEs in an ARD state in comparison with that of non-ARD individuals.

Earlier version: The relationship between DNA methylation and expression of RTEs with age was also investigated in our study. The methylation levels of all three RTE classes decrease with chronological age in our analysis, and the RTE expression is in inverse correlation with the methylation level in LINE and LTR but not in SINE. RTE sequences are repressed via DNA methylation and heterochromatinization40. In the aging process, surveillance of transcriptional regulators such as SIRT6 wanes, leading to hypomethylation and heterochromatin reduction4. L1 has been identified to be repressed through DNA methylation41. However, SINE has been suggested to be predominantly regulated by histone modification, and loss of DNA methylation has little effect on the transcription of SINEs33, as reflected in our analysis.

Revised version: The relationship between DNA methylation and expression of RTEs with age was also investigated in our study. The methylation levels of all three RTE classes decrease with chronological age in our analysis, and the RTE expression is in inverse correlation with the methylation level in LINE and LTR but not in SINE. RTE sequences are repressed via DNA methylation and heterochromatinization41. In the aging process, surveillance of transcriptional regulators such as SIRT6 wanes, leading to hypomethylation and heterochromatin reduction4. L1 has been identified to be repressed through DNA methylation42. However, SINE has been suggested to be predominantly regulated by histone modification, and loss of DNA methylation has little effect on the transcription of SINEs34, as reflected in our analysis.

Earlier version: Expression microarray has been successfully employed to study the expression of RTEs42–44. However, the probe design of microarrays does not accommodate specifically to the repeIIve and interspersed nature of RTEs45. To circumvent this issue, our analysis was conducted on the class and family levels but not subfamily levels of RTEs to retain enough probes (Supplementary Fig. 1) providing sufficient power for calculating the expression scores. We also supported our findings with scRNA-seq, particularly when investigating the association between RTE expression and chronological age, to circumvent the potential constraints posed by using microarray data. However, we are aware of the limitations imposed by using microarray in this study, particularly the low number of RTE probes in the expression microarray data. Our study can be enriched with the advent of large RNA-seq cohorts for aging studies in the future.

Revised version: Expression microarray has been successfully employed to study the expression of RTEs43-45. However, the probe design of microarrays does not accommodate specifically to the repeIIve and interspersed nature of RTEs46. To circumvent this issue, our analysis was conducted on the class and family levels but not subfamily levels of RTEs to retain enough probes (Supplementary Fig. 1) providing sufficient power for calculating the expression scores. We also supported our findings with scRNA-seq, particularly when investigating the association between RTE expression and chronological age, to circumvent the potential constraints posed by using microarray data. However, we are aware of the limitations imposed by using microarray in this study, particularly the low number of RTE probes in the expression microarray data. Although we included one recently published RNA-seq dataset to validate our microarray result, our study can be enriched with the advent of large RNA-seq cohorts for aging studies in the future.

Earlier version: Our exploration of multiple cell types based on a published PBMC scRNA-seq cohort unveiled plasma cell as the only cell type for which the expression of RTEs correlates with BAR signature scores. This matches a recent finding showing a late senescent-like phenotype marked by accumulation of TEs in plasma cells during pre-malignant stages46. Lastly, we explored a supercentenarian PBMC scRNA-seq cohort as a model for healthy aging and identified that RTE expression is generally decreased in supercentenarians. However, the BAR signature scores, particularly senescence and SASP scores, increased in the NK cells of supercentenarians compared to the normal-aged group, suggesting that immunosenescence is an important factor in aging.

Revised version: Our exploration of multiple cell types based on a published PBMC scRNA-seq cohort unveiled plasma cell as the only cell type for which the expression of RTEs correlates with BAR signature scores. This matches a recent finding showing a late senescent-like phenotype marked by accumulation of TEs in plasma cells during pre-malignant stages47. Lastly, we explored a supercentenarian PBMC scRNA-seq cohort as a model for healthy aging and identified that RTE expression is generally decreased in supercentenarians. However, the BAR signature scores, particularly senescence and SASP scores, increased in the NK cells of supercentenarians compared to the normal-aged group, suggesting that immunosenescence is an important factor in aging.

Removed in the revised version: Supporting information